# Homework: Python Exercises


## Name: <span style="color:blue"> Tyler Hudson </span>

## Utils

```python
# Extra imports for this lab that are beyond the scope of discussion
from typing import List, Dict, Tuple
import os
import gc
import traceback
import warnings
from pdb import set_trace

# Default seed
seed = 0
```

<!-- grade_id: cell-35888d71a1277e78, points: 0, grade: true, solution: false -->
```python
class TodoCheckFailed(Exception):
    pass

def todo_check(asserts, mute=False, **kwargs):
    locals().update(kwargs)
    failed_err = "You passed {}/{} and FAILED the following code checks:\n{}"
    failed = ""
    n_failed = 0
    for check, (condi, err) in enumerate(asserts):
        exc_failed = False
        if isinstance(condi, str):
            try:
                passed = eval(condi)
            except Exception:
                exc_failed = True
                n_failed += 1
                failed += f"\nCheck [{check+1}]: Failed to execute check [{check+1}] due to the following error...\n{traceback.format_exc()}"
        elif isinstance(condi, bool):
            passed = condi
        else:
            raise ValueError("asserts must be a list of strings or bools")

        if not exc_failed and not passed:
            n_failed += 1
            failed += f"\nCheck [{check+1}]: Failed\n\tTip: {err}\n"

    if len(failed) != 0:
        passed = len(asserts) - n_failed
        err = failed_err.format(passed, len(asserts), failed)
        raise TodoCheckFailed(err.format(failed))
    if not mute: print("Your code PASSED all the code checks!")

```

# Instructions

In this assignment, your core Python skills will be tested by having you complete or write various functions and classes using concepts such as lists, dictionaries, loops, and recursion. **No libraries are allowed to be used in this assignment!**

For each exercise, you will be given a set of instructions and a test function denoted with the prefix `TEST`. The testing function will contain a `todo_check()` function which will give you a rough estimate of whether your code is functioning as excepted. Feel free to write your own tests to make sure your code works with various different inputs. 

## Submission

1. Save the notebook.
2. Enter your name in the appropriate markdown cell provided at the top of the notebook.
3. Select `Kernel` -> `Restart Kernel and Run All Cells`. This will restart the kernel and run all cells. Make sure everything runs without errors and double-check the outputs are as you desire!
4. Submit the `.ipynb` notebook on Canvas.


# Python Exercises 

## List Operation (20 points)

Complete the `center_list()` function, which should take in a list `x` as input, and subtract the mean of `x` from every element in `x`. Be sure to store the mean of `x` in the variable `mean` and the centered list in the variable `centered_list`. Both `mean` and `centered_list` should be returned! Run the `TEST_center_list()` function to test your code.

**Example**

The output for the list `[5, 10, 15, 20, 25]` should be a mean of 15 and the new list `[-10.0, -5.0, 0.0, 5.0, 10.0]` 

**Hints**
- The built-in Python functions `sum()` and `len()` can be useful here.


<!-- grade_id: cell-de89530c0b47f42f, points: 0, grade: false, solution: true -->
```python
def center_list(x: List) -> Tuple[float, List]:
    mean: float = None
    centered_list: List = []
    # YOUR CODE HERE

    total = 0
    for i in x:
        total = total + i

    mean = total / len(x)

    for i in x:
        centered_list.append(i - mean)

    return mean, centered_list
```

<!-- grade_id: cell-2cdb8f7191d784c1, points: 20, grade: true, solution: false -->
```python
def TEST_center_list() -> None:
    print("{:=^50}".format('Inputs'))
    x = [5, 10, 15, 20, 25]
    print(f"x: {x}")
    
    print("{:=^50}".format('Outputs'))
    mean, centered_list = center_list(x)
    print(f"Mean: {mean}")
    print(f"Centered list: {centered_list}")
    
    todo_check([
        ('mean == 15.0', 'The expected output for `mean` is 15',), 
        ("centered_list ==  [-10.0, -5.0, 0.0, 5.0, 10.0]", 'The expected output for `centered_list` is [-10.0, -5.0, 0.0, 5.0, 10.0]')
    ],
    **locals())
    
TEST_center_list()
```

Output:
```
======================Inputs======================
x: [5, 10, 15, 20, 25]
=====================Outputs======================
Mean: 15.0
Centered list: [-10.0, -5.0, 0.0, 5.0, 10.0]
Your code PASSED all the code checks!
```

## Matrix Replacement (20 points)


Complete the `matrix_replacement()` function, which should take in a list of lists `X` (representing a matrix) as input. The goal of the `matrix_replacement()` function is to replace (i.e., filter) all `None` values with the integer value of 0. If you edited `X` directly or created a new matrix with the replaced values, be sure to return it. Run the `TEST_matrix_replacement()` function to test your code.

**Example**

The output for the list `[[1, None, 2], [None, 3, 1], [5, 6, None]]` should be `[[1, 0, 2], [0, 3, 1], [5, 6, 0]]`

**Hints**
- You will have to use nested for-loops to iterate over the matrix.


<!-- grade_id: cell-f055957a3bc60baa, points: 0, grade: false, solution: true -->
```python
def matrix_replacement(X: List[List]) -> List:
    # YOUR CODE HERE
    for i in range(len(X)):
        for j in range(len(X[i])):
            if X[i][j] == None:
                X[i][j] = 0
    
    return X
```

<!-- grade_id: cell-1191d977e0cfdf72, points: 20, grade: true, solution: false -->
```python
def TEST_matrix_replacement() -> None:
    print("{:=^50}".format('Inputs'))
    X = [[1, None, 2],[None, 3, 1],[5, 6, None]]
    print(f"x: {X}")
    
    print("{:=^50}".format('Outputs'))
    results = matrix_replacement(X)
    print(f"Results: {results}")
    
    todo_check([
        ('results == [[1, 0, 2], [0, 3, 1], [5, 6, 0]]', 'The expected output for `results` is [[1, 0, 2], [0, 3, 1], [5, 6, 0]]',), 
    ],
    **locals())
TEST_matrix_replacement()
```

Output:
```
======================Inputs======================
x: [[1, None, 2], [None, 3, 1], [5, 6, None]]
=====================Outputs======================
Results: [[1, 0, 2], [0, 3, 1], [5, 6, 0]]
Your code PASSED all the code checks!
```

## Histogram (20 points)
Complete the `matrix_replacement()` function, which should take in a list `x` as input. The goal of the `make_histogram()` function is to output a Python dictionary mapping each unique integer in `x` to the number of times it appears in `x`. Be sure to store the resulting histogram into the `results` variable and to return it. Run the `TEST_make_histogram()` function to test your code.

**Example**

The output for the list `[1, 2, 1, 3, 2, 1, 6, 2, 6, 1]` should be `{1: 4, 2: 3, 3: 1, 6: 2}`

**Hints**
- You will need to loop over `x` and check if the variable is already in the dictionary `results` to then add it or iterate the stored value.

<!-- grade_id: cell-47266446cd7e218b, points: 0, grade: false, solution: true -->
```python
def make_histogram(x: List) -> Dict:
    results: Dict = {}
    # YOUR CODE HERE

    for i in x:
        if i in results:
            results[i] += 1
        else:
            results[i] = 1
    
    return results
```

<!-- grade_id: cell-7399778738f23876, points: 20, grade: true, solution: false -->
```python
def TEST_make_histogram() -> None:
    print("{:=^50}".format('Inputs'))
    x = [1, 2, 1, 3, 2, 1, 6, 2, 6, 1]
    print(f"x: {x}")
    
    print("{:=^50}".format('Outputs'))
    results = make_histogram(x)
    print(f"Result: {results}")
    todo_check([
        ('isinstance(results, dict)', f'The result of `make_histrogram` must be of type dict, instead got `{type(results)}`'),
        ('results == {1:4, 2:3, 3:1, 6:2}', 'The expected output for `results` is {1:4, 2:3, 3:1, 6:2}',), 
    ],
    **locals())
TEST_make_histogram()
```

Output:
```
======================Inputs======================
x: [1, 2, 1, 3, 2, 1, 6, 2, 6, 1]
=====================Outputs======================
Result: {1: 4, 2: 3, 3: 1, 6: 2}
Your code PASSED all the code checks!
```

## Class and List Indexing (20 points)

Complete the `ListUtils()` class. To do so, follow the bellow instructions for implementing 6 different methods.

1. Implement the constructor method `__init__()`. This method should take a list `x` as input. Store `x` in the *class variable* `self.x`.
2. Implement a method called `list_length()` which takes no inputs but returns the length of the class variable `x`.
3. Implement a method called `first_n_elements()` which takes an integer `n` as input and returns the *first* `n` elements of the class variable `x`.
4. Implement a method called `last_n_elements()` which takes an integer `n` as input and returns the *last* `n` elements of the class variable `x`.
5. Implement a method called `index_n_elements()` which takes a list `indices` as input and returns the elements in the class variable `x` corresponding to the indices contained within`indices`.
6. Implement a method called `slice_list()` which takes two inputs: an integer `start` and an integer `end`. This method should return all elements between the start index `start` and up to the end index `end` (not include the end index) for the class variable `x`.


**Examples**

- For `list_length()`, the output should be 11.
- For `first_n_elements(n=5)`, the output should be `[1, 2, 3, 4, 5]`
- For `last_n_elements(n=5)`, the output should be `[7, 8, 9, 10, 11]`
- For `index_n_elements(indices=[0,4,7,10])`, the output should be `[1, 5, 8, 11]`
- For `slice_list(start=2, end=7)`, the output should be `[3, 4, 5, 6, 7]`

**Hints**
- Each method should take as input a reference to itself called `self` along with whatever other inputs are required by the instructions. 
- For most of the methods, you should only need to [slice](https://www.w3schools.com/python/python_strings_slicing.asp) the class variable `x`.
- If a method has multiple variables, define them in the order given in the instructions, otherwise your code might not pass the code checks.

<!-- grade_id: cell-24862a42e78e48f9, points: 0, grade: false, solution: true -->
```python
class ListUtils():
    # YOUR CODE HERE
    def __init__(self, x: List):
        self.x = x

    def list_length(self):
        return len(self.x)
    
    def first_n_elements(self, n: int):
        return self.x[:n]
    
    def last_n_elements(self, n: int):
        return self.x[-n:]
    
    def index_n_elements(self, n: List):
        listToReturn = []
        for i in n:
            listToReturn.append(self.x[i])

        return listToReturn
    
    def slice_list(self, start, end):
        return self.x[start:end]
        


```

<!-- grade_id: cell-580fc84702cc3961, points: 20, grade: true, solution: false -->
```python
def TEST_ListUtils() -> None:
    print("{:=^50}".format('Inputs'))
    x = list(range(1, 12))
    print(f"x: {x}")
    
    print("{:=^50}".format('Outputs'))
    list_utils = ListUtils(x)
    try:
        list_length = list_utils.list_length()
        print(f"list_length: {list_length}")
    except Exception:
        traceback.print_exc()
    try:
        first_n_elements = list_utils.first_n_elements(5)
        print(f"first_n_elements: {first_n_elements}")
    except Exception:
        traceback.print_exc()
    try:
        last_n_elements = list_utils.last_n_elements(5)
        print(f"last_n_elements: {last_n_elements}")
    except Exception:
        traceback.print_exc()    
    try:
        index_n_elements = list_utils.index_n_elements([0,4,7,10])
        print(f"index_n_elements: {index_n_elements}")
    except Exception:
        traceback.print_exc()
    try:
        slice_list = list_utils.slice_list(2,7)
        print(f"slice_list: {slice_list}")
    except Exception:
        traceback.print_exc()   

    todo_check([
        ('11 == list_utils.list_length()', 'The expected output for `list_utils.list_length()` is 11'),
        ('[1,2,3,4,5] == list_utils.first_n_elements(5)', 'The expected output for `list_utils.first_n_elements(5)` is [1, 2, 3, 4, 5]',), 
        ('[7,8,9,10,11] == list_utils.last_n_elements(5)', 'The expected output for `list_utils.last_n_elements(5)` is [7, 8, 9, 10, 11]',), 
        ('[1,5,8,11] == list_utils.index_n_elements([0,4,7,10])', 'The expected output for `list_utils.index_n_elements([0,4,7,10])` is [1, 5, 8, 11]',), 
        ('[3,4,5,6,7] == list_utils.slice_list(2,7)', 'The expected output for `list_utils.slice_list(2,7)` is [3, 4, 5, 6, 7]',), 
    ],
    **locals())

    
TEST_ListUtils()
```

Output:
```
======================Inputs======================
x: [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11]
=====================Outputs======================
list_length: 11
first_n_elements: [1, 2, 3, 4, 5]
last_n_elements: [7, 8, 9, 10, 11]
index_n_elements: [1, 5, 8, 11]
slice_list: [3, 4, 5, 6, 7]
Your code PASSED all the code checks!
```

## Recursion (20 points) 
Complete the `iterate_over_tree()` function, which should takes a list of lists `X` as input. The goal of the `iterate_over_tree()` function is to return the number of nodes `nodes` and the leaf nodes values `leafs` in a tree. The tree structure is defined using a list of lists such as `[root, [subtree_1 [leaf1_], [leaf_2]] , [leaf_3] , [subtree_2 [leaf_4]]]` where a leaf node is indicated by a list with the length of 1. To solve this problem, you will need to use [recursion](https://realpython.com/python-recursion/) with `iterate_over_tree()` function. 

**Examples**
- Below is an expanded version of the list given in the instructions as an example to better visualize how the list represents a tree.
```Python
[root, 
    [subtree_1 
        [leaf1_], 
        [leaf_2],
    ],
    [leaf_3],
    [subtree_2,
         [leaf_4]
    ]
]
```
- The output for `[1, [2, [3], [4]], [5], [6, [7, [8]], [9]]]` should be 9 nodes with the leafs `[3, 4, 5, 8, 9]`

**Hints**

- You might need to use Python's `isinstance()` built-in function to check if an element is a list, if it is, then you can initiate recursion.
- Remember the `len()` built-in function is used to check the length of a list.

<!-- grade_id: cell-ce2a9e122bc11b69, points: 0, grade: false, solution: true -->
```python
def iterate_over_tree(X: List[List]) -> Tuple[int, List]:
    nodes: int = 0
    leafs: List = []
    # YOUR CODE HERE

    for i in X:
        if isinstance(i, list) and len(i) > 1:
            child_nodes, child_leafs = iterate_over_tree(i)
            nodes = nodes + child_nodes 
            leafs.extend(child_leafs) 
        else:
            nodes = nodes + 1
            if isinstance(i, list) and len(i) == 1:
                leafs.append(i[0])

    return nodes, leafs

print(iterate_over_tree([1, [2, [3], [4]], [5], [6, [7, [8]], [9]]]))
```

Output:
```
(9, [3, 4, 5, 8, 9])
```

<!-- grade_id: cell-40556244fba60849, points: 20, grade: true, solution: false -->
```python
def TEST_iterate_over_tree():
    print("{:=^50}".format('Inputs'))
    X = [1, [2, [3], [4]], [5], [6, [7, [8]], [9]]]
    print(f"X: {X}")

    print("{:=^50}".format('Outputs'))
    count, leafs = iterate_over_tree(X)
    print(f"Count: {count}")
    print(f"leafs: {leafs}")

    todo_check([
        ('count == 9', 'The expected output for `count` is 9',),
        ('leafs == [3, 4, 5, 8, 9]', 'The expected output for `leafs` is [3, 4, 5, 8, 9]',), 
    ],
    **locals())
TEST_iterate_over_tree()
```

Output:
```
======================Inputs======================
X: [1, [2, [3], [4]], [5], [6, [7, [8]], [9]]]
=====================Outputs======================
Count: 9
leafs: [3, 4, 5, 8, 9]
Your code PASSED all the code checks!
```